![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

```python
import numpy as np
import matplotlib.pyplot as plt
```

# Questão 1: Bisseção ao infinito

Uma outra forma de encontrar raizes de $x e^x = y$ consiste em fazer uma bisseção "até o infinito".
Vejamos como isso se aplica no caso de $+\infty$.

Como há um número **finito** de números de ponto flutuante que seu computador conhece,
podemos usar uma progressão geométrica para tentar chegar no infinito.

Se no intervalo $[a, a\cdot r]$ a função não trocar de sinal, podemos tentar o intervalo
$[a\cdot r, a\cdot r^2]$, e depois o intervalo $[a\cdot r^2, a\cdot r^3]$ e assim sucessivamente.

## Questão 1.1: PGs ao infinito

Escreva uma função que calcula o comprimento da PG $a r^n$ até chegar no infinito.

<!-- grade_id: cell-c4d8baa587bac6d4, points: 0, grade: false, solution: true -->
```python
def comprimento_pg(a, r):
    """Comprimento da PG de razão  r  começando em  a,  até chegar no +infinito de ponto flutuante"""
    assert a > 0 and r > 1
    a, r = float(a), float(r)
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d738818c07c00d8c, points: 1, grade: true, solution: false -->
```python
assert 1000 < comprimento_pg(1, 2.) < 2000
```

<!-- grade_id: cell-177b86067a2baf3b, points: 1, grade: true, solution: false -->
```python
assert 250 < comprimento_pg(1, 10.) < 350
```

<!-- grade_id: cell-2b3807501148637a, points: 1, grade: true, solution: false -->
```python
assert 250 < comprimento_pg(1e10, 10.) < 350
```

## Questão 1.2: Máximo de bisseções

Quantas bisseções é possível fazer em um intervalo $[a,b]$ até chegar no limite de precisão do computador?

<!-- grade_id: cell-d1cc08dbefd29ef2, points: 0, grade: false, solution: true -->
```python
def bissect_max(a, b):
    """Conta o número de bisseções no intervalo  [a, b]  até chegar na precisão de representação."""
    nbiss = 0
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f1d46061b1a36689, points: 1, grade: true, solution: false -->
```python
assert 50 <= bissect_max(1, 2) <= 56
assert 50 <= bissect_max(1, 4) <= 56
```

<!-- grade_id: cell-f2a4944177470fbb, points: 1, grade: true, solution: false -->
```python
assert bissect_max(1,8) >= bissect_max(1,4)
assert bissect_max(1,8) >= bissect_max(2,8)
```

<!-- grade_id: cell-ac8baba898217b6f, points: 1, grade: true, solution: false -->
```python
assert bissect_max(1,8) >= bissect_max(3,8)
```

Porque é sempre verdade que `bissect_max(a,b) >= bissect_max(c,d)` se $[c,d] \subset [a,b]$?

<!-- grade_id: cell-93416178f2ffc8aa, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 1.3: Calibrando duas bisseções

Qual razão $r$ da PG você usaria para encontrar uma raiz positiva,
levando em conta
- o número de bisseções para encontrar o intervalo; e
- o número de bisseções para encontrar a raiz dentro do intervalo?

Explique sua ideia na caixa abaixo, e faça as contas necessárias na caixa seguinte.

<!-- grade_id: cell-cdc5d304cb874a8b, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-7b2b2ed147be726b, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Questão 1.4: "Fase 0"

Agora, implemente uma função que encontra um intervalo onde há uma raiz de $f(x) = y$,
sabendo que há uma raiz no intervalo $[1, +\infty)$.

<!-- grade_id: cell-041d37d1e07388cd, points: 0, grade: false, solution: true -->
```python
def fase0(f, y):
    """Retorna um intervalo onde há uma raiz de  f(x) = y."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e4cdf4206efc804f, points: 1, grade: true, solution: false -->
```python
def sqr(x):
    return x**2

a, b = fase0(sqr, 10)
assert sqr(a) <= 10 <= sqr(b) or sqr(b) <= 10 <= sqr(a)
```

<!-- grade_id: cell-5fde738ed830b7bf, points: 1, grade: true, solution: false -->
```python
def slow_pot(x):
    return x**0.1

a, b = fase0(slow_pot, 10)
assert slow_pot(a) <= 10 <= slow_pot(b) or slow_pot(b) <= 10 <= slow_pot(a)
```

<!-- grade_id: cell-31d91f4937ad3e8a, points: 1, grade: true, solution: false -->
```python
def cauchy(x):
    return 1/(1 + x**2)

a, b = fase0(cauchy, 1e-10)
assert cauchy(a) <= 1e-10 <= cauchy(b) or cauchy(b) <= 1e-10 <= cauchy(a)
```

Agora, modifique a `fase0` para receber valores de $a$ e $r$ quaisquer.
Isso permitirá começar a busca tanto para valores positivos como negativos,
e também "no infinito" ou "perto de zero".

<!-- grade_id: cell-58025fe0b66b2396, points: 0, grade: false, solution: true -->
```python
def fase0_geral(f, y, a, r):
    """Retorna um intervalo contendo uma raiz de  f(x) = y,  usando a PG de termo inicial  a  e razão  r."""
    assert a != 0
    assert r > 0 and abs(r) != 1
    # YOUR CODE HERE
    raise NotImplementedError()
```

Explique os `assert`s iniciais.

<!-- grade_id: cell-5b71264c4743cbf5, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-945fc3bf2f130ff1, points: 1, grade: true, solution: false -->
```python
a, b = fase0_geral(sqr, 10, -1, 2)
assert sqr(a) <= 10 <= sqr(b) or sqr(b) <= 10 <= sqr(a)
```

<!-- grade_id: cell-287e669d1e97c5a5, points: 1, grade: true, solution: false -->
```python
def xlogx(x):
    return x * np.log(x)

a, b = fase0_geral(xlogx, -1e-20, 0.5, 0.1)
assert xlogx(a) <= -1e-20 <= xlogx(b) or xlogx(b) <= -1e-20 <= xlogx(a)
```

## Questão 1.5: Bisseção final

Modifique o código da Bisseção para encontrar a raiz de $f(x) = y$, até a precisão do computador,
a partir dos termos de uma PG.

Retorne, também, o número de pontos onde você terá avaliado $f$ - isso vai necessitar alterar a função de fase 0.

<!-- grade_id: cell-a1349ad9bf427a3c, points: 0, grade: false, solution: true -->
```python
def fase0_geral(f, y, a, r):
    """Retorna um intervalo contendo uma raiz de  f(x) = y,  usando a PG de termo inicial  a  e razão  r."""
    assert a != 0
    assert r > 0 and abs(r) != 1
    # YOUR CODE HERE
    raise NotImplementedError()
    
def bisseção_PG(f, y, a, r):
    """Solução até a precisão da máquina para  f(x) = y."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a47d44faa3f14dd6, points: 2, grade: true, solution: false -->
```python
r, nevals = bisseção_PG(sqr, 1000, 1, 2)
assert abs(sqr(r) - 1000) < 1e-13
assert 50 <= nevals <= 100
```

<!-- grade_id: cell-c473465dda9ff5e7, points: 1, grade: true, solution: false -->
```python
r, nevals = bisseção_PG(xlogx, 1000, 1, 3)
assert abs(xlogx(r) - 1000) < 1e-3
assert 50 <= nevals <= 100
```

<!-- grade_id: cell-913e3c644f4f2ae3, points: 1, grade: true, solution: false -->
```python
def xex(x):
    return x * np.exp(x)

r, nevals = bisseção_PG(xex, -1e-30, -1, 10)
assert abs(xex(r) + 1e-30) < 1e-40
assert 50 <= nevals <= 70
```

<!-- grade_id: cell-b1631c7326b9eb7a, points: 1, grade: true, solution: false -->
```python
r, nevals = bisseção_PG(xlogx, -1e-20, 0.5, 0.1)
assert abs(xlogx(r) + 1e-20) < 1e-30
assert 50 <= nevals <= 100
```

Porque a tolerância dos últimos testes é tão pequena?

<!-- grade_id: cell-45e0e0366ba0a5de, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 1.6: Generalizando ainda mais

Como você adaptaria a busca de um intervalo para encontrar uma raiz de $f(x) = y$ para funções que não são monótonas?

Algumas ideias para pensar:

- $x \cos(x)$
- $x + \sin(x)$

<!-- grade_id: cell-fdacb4aa36b4346c, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

# Questão 2: O Método da Secante

Uma variante do método de Newton, que dispensa o uso da derivada,
é o método da secante.

A ideia é simples: em vez de usar a tangente, dada pela derivada $f'(x)$,
o método da secante usa 2 pontos consecutivos $x_0$ e $x_1$ para calcular uma **secante** ao gráfico de $f$,
e com isto determinar uma interseção desta reta com o eixo dos $x$.
Esta interseção é chamada de $x_2$, e o método prossegue construindo $x_3$ a partir de $x_1$ e $x_2$
e assim por diante.

## Questão 2.1: Equações

Dê a equação da reta secante a $f$, passando pelos pontos $x_0$ e $x_1$.

<!-- grade_id: cell-1eedd3eb4d36b88f, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Deduza a equação do "passo da secante", que a partir $x_0$ e $x_1$ calcula a raiz da secante ao gráfico de $f$.

<!-- grade_id: cell-89fe6d21bb0d4154, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

O método da secante é inspirado no método de Newton.
Explique porque ele também pode entrar em _loop_ infinito.

<!-- grade_id: cell-8553d8123bdac585, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 2.2: Implementação

Escreva o método da secante para encontrar uma raiz de $f$.

Retorne todos os pontos visitados pelo método.

<!-- grade_id: cell-c0fbfc192b4aedfb, points: 0, grade: false, solution: true -->
```python
def secante(f, x0, x1, xtol=1e-8, ytol=1e-8, maxiter=100):
    """Encontra uma raiz de  f,  começando pelos pontos  x0 e x1,  pelo método da secante."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b2d65d0ef437e218, points: 2, grade: true, solution: false -->
```python
def f2(x):
    return x**2 - 2

pts = secante(f2, 0, 1)
assert 5 <= len(pts) <= 10
assert abs(f2(pts[-1])) <= 1e-5
```

<!-- grade_id: cell-10e8cfbc06341a84, points: 2, grade: true, solution: false -->
```python
def f3(x):
    return x**3 - 3

pts = secante(f3, 0, 1)

assert 8 <= len(pts) <= 12
assert abs(f3(pts[-1])) <= 1e-5
```

## Questão 2.3: Gráficos

Faça o gráfico da distância (em valor absoluto)
entre os pontos produzidos pelo método da secante e a raiz verdadeira,
para as funções $f2$ e $f3$ acima.

<!-- grade_id: cell-1a8126f735c32bfe, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-282f1138902672b7, points: 2, grade: true, solution: false -->
```python
assert len(ax.lines) == 2
assert len(ax.get_legend().texts) == 2
```

De que forma(s) este gráfico parece com o equivalente para o Método de Newton?

<!-- grade_id: cell-ee7f5cb912169859, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 2.4: Generalização

Agora, transforme o seu código para encontrar uma raiz de $f(x) = y$.

<!-- grade_id: cell-37ecabe4fbd2ab4d, points: 0, grade: false, solution: true -->
```python
def secante_inv(f, y, x0, x1, xtol=1e-8, ytol=1e-8, maxiter=100):
    """Encontra uma solução de  f(x) = y,  começando pelos pontos  x0 e x1,  pelo método da secante."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2a826bf7700bb0c1, points: 2, grade: true, solution: false -->
```python
pts = secante_inv(sqr, 2, 0, 1)
pts_more = secante_inv(sqr, 2, 0, 1, xtol=1e-15, ytol=1e-15)

assert len(pts) < len(pts_more)
assert abs(pts[-1]**2 - 2) > abs(pts_more[-1]**2 - 2)
```

Explique porque estes dois testes fazem sentido.

<!-- grade_id: cell-d68950348b9d12ac, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-d46c38019a769e5c, points: 1, grade: true, solution: false -->
```python
assert len(pts_more) == len(pts) + 1
```

Explique porque este teste mostra a velocidade do método da secante.

<!-- grade_id: cell-d3d4b8bf63db3bdb, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 2.5: Mais gráficos!

Para vários valores de $y$, faça em um mesmo eixo o gráfico da distância até a raiz
para o método da secante resolver $x^2 = y$, começando de $(0, 1)$.

Em particular, encontre valores de $y$ que façam que o método da secante demore (bem) mais do que para $y = 2$.

<!-- grade_id: cell-a638b886004599a3, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
```

<!-- grade_id: cell-a1a0f56083060410, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) >= 5
```

<!-- grade_id: cell-4eab8247581ab2cd, points: 1, grade: true, solution: false -->
```python
assert len(ax.get_legend().texts) >= 5
```

<!-- grade_id: cell-09546580c7794f3d, points: 1, grade: true, solution: false -->
```python
xmax = max(max(l.get_xdata()) for l in ax.lines)
assert xmax >= 18
```

Explique o comportamento que observado.

<!-- grade_id: cell-07cedca930738508, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 2.6: Bugs finais

Observe o código abaixo:

```python
secante_inv(xex, 100, 0, 1)
```

Porquê o método da secante terminou longe demais de uma raiz?

<!-- grade_id: cell-220cfe43327adf0f, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

O que poderia ser feito para evitar este tipo de situação?

<!-- grade_id: cell-a487e6248cbf174e, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE