![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

# Integrais e Primitivas

Neste teste, vamos usar os métodos de integração para calcular primitivas de funções.

De forma análoga ao caso de EDOs, vamos retornar dois `np.array`s:
- um com os pontos onde a primitiva foi calculada; e
- outro com o valor da aproximação numérica da primitiva.

```python
import numpy as np
import matplotlib.pyplot as plt
```

# Questão 0: o método do retângulo

Para evitar efeitos de bordo e arredondamento, vamos usar o método do retângulo criando pontos com `np.linspace`:

<!-- grade_id: cell-e58721f1515c51de, points: 0, grade: false, solution: true -->
```python
def retangulo(f, a, b, n=1000):
    if a == b:
        return 0
    pts, h = np.linspace(a, b, num=n, endpoint=False, retstep=True)
    # YOUR CODE HERE
    raise NotImplementedError()
```

Explique porque os pontos criados têm `endpoint=False`

<!-- grade_id: cell-1549d22006ecb487, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-a91970e14a540c4c, points: 1, grade: true, solution: false -->
```python
ans = (np.cos(0) - np.cos(1))
assert abs(retangulo(np.sin, 0, 1) - ans) < 1e-3
```

<!-- grade_id: cell-4096352942140c09, points: 1, grade: true, solution: false -->
```python
ans = (np.cos(0) - np.cos(1))
assert abs(retangulo(np.sin, 0, 1) - ans) > 1e-4
```

<!-- grade_id: cell-a76174554c52c73d, points: 1, grade: true, solution: false -->
```python
ans = (np.cos(0) - np.cos(2))
assert 1e-5 < abs(retangulo(np.sin, 0, 2, n=10**4) - ans) < 1e-4
```

# Questão 1: O jeito fácil

Escreva uma função `primitiva_retangulo`,
que calcula a integral de $a$ até cada um dos pontos escolhidos,
**incluindo** $a$,
chamando o método do retângulo para cada um deles.

<!-- grade_id: cell-ca2f7a5c2102974b, points: 0, grade: false, solution: true -->
```python
def primitiva_retangulo(f, a, b, h=0.01):
    """Primitiva da função  f, usando um passo de tamanho  h  no intervalo  [a,b]."""
    pts = np.arange(a, b, h)
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f782ccb631628bac, points: 1, grade: true, solution: false -->
```python
pts, Fs = primitiva_retangulo(np.sin, 0, 1)
assert len(Fs) == len(pts)
```

<!-- grade_id: cell-610d0a8c1b3b1ee5, points: 1, grade: true, solution: false -->
```python
assert np.allclose(Fs, 1 - np.cos(pts), atol=1e-2)
```

<!-- grade_id: cell-f08c296c1f556aa0, points: 1, grade: true, solution: false -->
```python
assert not np.allclose(Fs, 1 - np.cos(pts), atol=1e-3)
```

## Questão 2: Gráficos

Faça os gráficos dos erros de cálculo das primitivas, no intervalo $[0, 1.5]$, das seguintes funções:
- Seno
- Exponencial
- Cosseno

para quatro valores de passo: `1e-2`, `5e-3`, `2e-3` e `1e-3`.

Use 3 eixos, um para cada função, e 4 curvas em cada eixo.

<!-- grade_id: cell-893a01adad1a224e, points: 1, grade: true, solution: true -->
```python
fig, axs = plt.subplots(ncols=3, figsize=(18,4))
hs = [1e-2, 5e-3, 2e-3, 1e-3]

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-63283f9cb6e27e97, points: 2, grade: true, solution: false -->
```python
for ax in axs:
    assert len(ax.lines) == 4
    assert len(ax.get_legend().texts) == 4
    assert len(ax.get_title()) > 0
```

<!-- grade_id: cell-b8348c3801bc1956, points: 1, grade: true, solution: false -->
```python
for ax in axs:
    for l in ax.lines:
        assert max(np.abs(l.get_ydata())) < 2e-2
```

Qual o erro, na extremidade final do intervalo, para cada um dos métodos e valores de $h$?

Imprima uma tabela, com 3 colunas e 4 linhas (fora as "de título"), usando `print` formatado.

<!-- grade_id: cell-85eff91dd1f86421, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Como os erros se comportam ao diminuir o passo?  Isso é esperado?

<!-- grade_id: cell-2de92a806cae4abd, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Como os erros se comportam ao longo do eixo dos $x$, ou seja, ao calcular primitivas "cada vez mais longe"?  Como você explica este fenômeno?

<!-- grade_id: cell-b016645ddc678451, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Como os erros se comportam para as diferentes funções?  Você esperava isso?

<!-- grade_id: cell-0842df6fbab5c7a4, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 3: O jeito mais rápido

Na questão anterior, ao calcular a primitiva de $[a, a+2h]$, era necessário calcular $f(a)$ e $f(a+h)$.
Ora, $f(a)$ já tinha sido calculada para a primitiva em $[a, a+h]$ (era o único termo!),
então podemos evitar estas contas repetidas.

Vamos fazer uma nova primitiva, `primitiva_rapida`, que calcula apenas uma vez `f(pts)`,
e calcula as somas parciais usando [`np.cumsum`](https://numpy.org/doc/stable/reference/generated/numpy.cumsum.html).

Explique porque é ruim (neste caso!) que `np.cumsum` **não** retorne a soma vazia.

```python
np.cumsum([1,2])
```

<!-- grade_id: cell-c5a49030557d7f4c, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Para incluir a soma vazia (que vale zero!), use `np.hstack`.

```python
np.hstack([[0], [1, 3]])
```

<!-- grade_id: cell-66631c20ed24bf35, points: 0, grade: false, solution: true -->
```python
def primitiva_rapida(f, a, b, h=0.01):
    """Primitiva da função  f, usando um passo de tamanho  h  no intervalo  [a,b]."""
    pts = np.arange(a, b, h)
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-50f5e7584dce3ef8, points: 1, grade: true, solution: false -->
```python
pts, Fs = primitiva_rapida(np.sin, 0, 1)
assert len(pts) == len(Fs)
```

<!-- grade_id: cell-3df3bdeb658b648d, points: 1, grade: true, solution: false -->
```python
assert np.allclose(Fs, 1 - np.cos(pts), atol=1e-2)
```

<!-- grade_id: cell-4d749fbf56409d26, points: 1, grade: true, solution: false -->
```python
assert not np.allclose(Fs, 1 - np.cos(pts), atol=1e-3)
```

Ambas as funções devem ser (a menos de erro de truncamento / arredondamento) iguais:

<!-- grade_id: cell-8e62b317e12f9bd6, points: 2, grade: true, solution: false -->
```python
_, Fs_original = primitiva_retangulo(np.sin, 0, 1)
assert np.allclose(Fs, Fs_original, atol=1e-15, rtol=1e-15)
```

## Questão 4: E terá valido a pena?

Observe as duas caixas abaixo

```python
%time primitiva_retangulo(np.sin, 0, 1, h=1e-4)
%time primitiva_retangulo(np.sin, 0, 1, h=1e-3)
None
```

```python
%time primitiva_rapida(np.sin, 0, 1, h=1e-4)
%time primitiva_rapida(np.sin, 0, 1, h=1e-3)
None
```

Quão mais rápido é o método usando `np.cumsum`?

<!-- grade_id: cell-ffce7e6028d2340e, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Agora, observe com mais detalhe:

```python
%timeit primitiva_retangulo(np.sin, 0, 1, h=1e-4)
%timeit primitiva_retangulo(np.sin, 0, 1, h=1e-3)
None
```

```python
%timeit primitiva_rapida(np.sin, 0, 1, h=1e-4)
%timeit primitiva_rapida(np.sin, 0, 1, h=1e-3)
None
```

Quantas vezes mais lento é o método "simples" com passo 10 vezes menor?

E quantas vezes mais lento é o método `np.cumsum` com passo 10 vezes menor?

Isso ajuda a explicar as diferenças acima?

<!-- grade_id: cell-af8d633f85cf0dbe, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE