![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

```python
import numpy as np
```

# Teste 1: Bisseção

## Questão 1: Número de iterações da bisseção, e de chamadas à função `f`

Generalize o algoritmo da bisseção _iterativa_ para retornar,
- a (aproximação da) raiz,
- o número de bisseções (divisões por 2 do intervalo) feitas,
- o número de vezes que você chamou a função `f`

Use como critério de parada
- `xtol`, e
- `maxiter`,

ou seja, o algoritmo termina quando:
- seja possível garantir que o erro (absoluto) da resposta ("em $x$") seja menor do que `xtol`; OU
- quando o algoritmo já tiver feito `maxiter` itererações.

<!-- grade_id: cell-9c2f0bf1fe48da91, points: 0, grade: false, solution: true -->
```python
def bissecao(f, a, b, xtol=1e-8, maxiter=10):
    # YOUR CODE HERE
    raise NotImplementedError()
```

### 1.1: Testando com $\sqrt2$:

```python
def p2(x): return x**2 - 2
```

<!-- grade_id: cell-6216e0119f2a0e78, points: 1, grade: true, solution: false -->
```python
raiz, niters, ncalls = bissecao(p2, 0, 2, maxiter=100)
assert abs(raiz - np.sqrt(2)) < 1e-8
assert 25 <= niters <= 30
assert 25 <= ncalls
```

<!-- grade_id: cell-caee585cfc6b637f, points: 1, grade: true, solution: false -->
```python
raiz, niters, ncalls = bissecao(p2, 0, 2, xtol=1e-3)
assert abs(raiz - np.sqrt(2)) < 2e-3
assert 9 <= niters <= 11
```

<!-- grade_id: cell-ed17d59cca42278d, points: 1, grade: true, solution: false -->
```python
raiz, niters, ncalls = bissecao(p2, 0, 2)
assert abs(raiz**2 - 2) > 1e-3
assert niters == 10
```

<!-- grade_id: cell-19379484bc5a3adf, points: 1, grade: true, solution: false -->
```python

```

### 1.2: Tolerância, número de iterações, respostas, ...

Observe o seguinte código:

```python
for err in np.logspace(-8, 0, num=9):
    r, ni, nc = bissecao(p2, 0, 2, xtol=err)
    print(f"xtol = {err:.1e} --> raiz = {r:.10f}, #iter = {ni:3d}")
```

Comente

<!-- grade_id: cell-856d87a71e8f8ec5, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Como seria o equivalente se o critério fosse o número de iterações?

<!-- grade_id: cell-c6d45f6ca0f539e4, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Comente, também.

<!-- grade_id: cell-b7931432898e9741, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### 1.3: Agora, um polinômio um pouco mais complicado:

```python
def p4(x): return x**4 - 3*x + 1
```

<!-- grade_id: cell-b8f933e77c525f69, points: 1, grade: true, solution: false -->
```python
r1, n1, c1 = bissecao(p4, 0, 1, maxiter=100)
assert p4(r1) < 1e-6
```

<!-- grade_id: cell-fab8ee9c8e55f544, points: 1, grade: true, solution: false -->
```python
r2, n2, c2 = bissecao(p4, 1, 2, maxiter=100)
assert n1 == n2
assert p4(r2) < 1e-6
```

<!-- grade_id: cell-faf0c1fe3a384747, points: 1, grade: true, solution: false -->
```python
r3, n3, c2 = bissecao(p4, 1, 3, maxiter=100)
assert n3 == n2+1
assert r3 == r2
```

<!-- grade_id: cell-1ec4a98add6ac67a, points: 1, grade: true, solution: false -->
```python

```

### 1.4: Relação entre `ncalls` e `niters`

Muito provavelmente, é possível calcular `ncalls` a partir de `niters`.
Qual a relação, para o seu programa, entre estes valores?
Porquê?

<!-- grade_id: cell-4a9c83444fc8ee6a, points: 1, grade: true, solution: true -->
```python
# Use esta caixa para evidenciar a relação, e a de baixo para explicar qual é, e porquê
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7d3444498035a363, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 2: Bisseção em listas ordenadas

Se desejamos encontrar um número em uma lista, podemos percorrer cada um de seus índices
até achar o número, e retornamos o índice correspondente.

Mas, se a lista `l` possui números em ordem,
é possível usar uma variante da bisseção para encontrar este índice.
Observe que agora desejamos retornar um número **inteiro** (o índice!),
e que índices para listas sempre devem ser inteiros, portanto:
- cuidado ao escolher o "ponto médio"
- mas o critério de parada é "quando achar o índice".

### 2.1: Encontrando um número

Implemente a função `bissect_list(l, v)`, que retorna
- o índice `k` do elemento na lista `l`, que vale `v`; e
- o número de acessos à lista `l`.

Suponha que a lista `l` é não-decrescente (ou seja, `l[i] <= l[i+1]`).

<!-- grade_id: cell-8c9f49b0ba8b5d04, points: 0, grade: false, solution: true -->
```python
def bissect_list(l, v):
    """Índice  k  na lista  l , crescente, tal que l[k] = v.  None caso não exista."""
    n = len(l)
    naccess = 0
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d0e135df902b3955, points: 1, grade: true, solution: false -->
```python
assert bissect_list([], 42) == (None,0)
```

<!-- grade_id: cell-aa0a59b135a65159, points: 1, grade: true, solution: false -->
```python
l = [4]
assert bissect_list(l, 4) == (0,1)
```

<!-- grade_id: cell-ba03109833002639, points: 1, grade: true, solution: false -->
```python
assert bissect_list(l, 5)[0] is None
```

<!-- grade_id: cell-84b5ba3a3384d223, points: 1, grade: true, solution: false -->
```python
l = [1,3,6,10,15,21,28]
for i,v in enumerate(l):
    idx, n = bissect_list(l, v)
    assert idx == i, f"v={v}, i={i}, idx={idx}"
    print(n)
```

Qual o maior número de acessos à lista que você precisou para encontrar um valor nela?
Este número depende da posição na lista?  Explique!

<!-- grade_id: cell-9059e7d02c8a0881, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### 2.2: Uma lista muito maior

```python
l = np.arange(-4000, 6000)
np.random.seed(1)
valores_testes = np.random.randint(low=-4000, high=6000, size=10)
for v in valores_testes:
    idx, n = bissect_list(l, v)
    assert l[idx] == v
    assert n <= 20
    print(f"v={v:5d}, idx={idx:5d}, n={n:3d}")
```

Como o número de acessos varia, nesta outra lista?
E como ele se compara com o tamanho da lista?

<!-- grade_id: cell-1853e505f297b9b5, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### 3.3: Listas crescentes e decrescentes.

Altere o código anterior, para receber um argumento a mais, `decr`, que indica que a lista é **decrescente**.
Não reordene a lista!

<!-- grade_id: cell-e442fe23480d4b8d, points: 0, grade: false, solution: true -->
```python
def bissect(l, v, decr=False):
    """Index  k  on a monotonic list  l  such that  l[k] == v.
    Returns  None  if  not present.
    If  l  is decresing, pass  decr=True."""
    n = len(l)
    naccess = 0
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8d0e1c185171dbce, points: 1, grade: true, solution: false -->
```python
l = [1,3,6,10,15,21,28,35,36]
assert bissect(l, 15)[0] == 4
```

<!-- grade_id: cell-75e5e557da5782e9, points: 1, grade: true, solution: false -->
```python
l_rev = [36, 35, 28, 21, 15, 10, 6, 3, 1]
assert bissect(l_rev, 15, decr=True)[0] == 4
```

<!-- grade_id: cell-bd452b4939210a9a, points: 1, grade: true, solution: false -->
```python
l_rev = [36, 35, 28, 21, 15, 10, 6, 3, 1]
assert bissect(l_rev, 16, decr=True)[0] == None
```

<!-- grade_id: cell-4f806045de0b056c, points: 1, grade: true, solution: false -->
```python
l3 = [36, 35, 28, 21, 15, 10, 9, 6, 3, 2, 1]
assert bissect(l3, 15, decr=True)[0] == 4
```

### 2.4: Bônus

O número de acessos à lista para encontrar um número muda se esta está invertida?
Explique!

<!-- grade_id: cell-def3480b5f0dcea3, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE