![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

# Integrando funções menos regulares

```python
import numpy as np
import matplotlib.pyplot as plt
```

<!-- grade_id: cell-25b2a1e4698b3883, points: 0, grade: false, solution: true -->
```python
def midpoint(f, a, b, n=1000):
    # YOUR CODE HERE
    raise NotImplementedError()

def simpson(f, a, b, n=1000):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
def f(x):
    return np.abs(np.sin(x))
```

## Questão 1: Primitiva

Dê uma primitiva analítica de $f$.

Dica: use a função `np.divmod` (versão vetorizada da `divmod` do python).

<!-- grade_id: cell-ab540a7e7e80c2e9, points: 0, grade: false, solution: true -->
```python
def F(x):
    """Primitiva de |sin(x)|, exata até a precisão do computador."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3a001c03b90f0783, points: 1, grade: true, solution: false -->
```python
assert np.isclose(F(10+0.1), F(10) + 0.1 * f(10), atol=1e-2)
```

<!-- grade_id: cell-324c73c9738ddcb5, points: 1, grade: true, solution: false -->
```python
assert np.isclose(F(2+0.01), F(2) + 0.01 * f(2), atol=1e-4)
```

Explique as tolerâncias dos testes acima.

<!-- grade_id: cell-705f645cd463d9a7, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-7e35b488ce517af7, points: 1, grade: true, solution: false -->
```python
ts = np.linspace(0,10, num=200)
Fts = F(ts)
assert np.all(Fts[1:] > Fts[:-1])
```

## Questão 2: Erros

Faça um gráfico dos erros de integração de $f$
pelos métodos do ponto médio e Simpson, no intervalo $[0,3]$

<!-- grade_id: cell-f522838820185ccf, points: 1, grade: true, solution: true -->
```python
a, b = 0, 3
exact = F(b) - F(a)
ns = np.logspace(3,13, num=20, dtype=int, base=2)

# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
```

<!-- grade_id: cell-9160718d1a3426a9, points: 2, grade: true, solution: false -->
```python
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2
```

```python
ax = None
```

O que você observa neste gráfico?

<!-- grade_id: cell-e098c0e743ef6a81, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Agora, repita o gráfico para o intervalo $[0,6]$:

<!-- grade_id: cell-e771e7f45189eb47, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
```

<!-- grade_id: cell-ffe3f5d3cdbd5614, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2
```

<!-- grade_id: cell-bed176376358bb71, points: 1, grade: true, solution: false -->
```python
for l in ax.lines:
    assert 1e-10 < min(l.get_ydata()) < 1e-7
```

```python
ax = None
```

Como esse gráfico é diferente do outro?

<!-- grade_id: cell-d6d2104f2e3363f0, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 3: Encontrando o intervalo contendo $\pi$

Como $|\sin(x)|$ não é derivável em $\pi$, vamos ver quanto erro tem no sub-intervalo que "atravessa" $\pi$.

Escreva a função `I_x(a, b, x, n)` que retorna o intervalo contendo $x$ ao dividir $[a,b]$ em $n$ partes.
Você pode assumir que x nunca vai ser uma das extremidades $a_k$.

<!-- grade_id: cell-ebebd4ecbff8d451, points: 0, grade: false, solution: true -->
```python
def I_x(a, b, x, n):
    """The interval  [xi, xi+h]  containing  x, when  [a,b]  is divided in  n  equal parts.
    Returns  None  if  x  is outside  [a,b]"""
    if (x-a)*(x-b) > 0:
        return
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6a808a637616f751, points: 1, grade: true, solution: false -->
```python
l, r = I_x(0, 6, np.pi, 10)
assert np.isclose(l, 3)
assert np.isclose(r, 3.6)
```

<!-- grade_id: cell-046e03f8c340519a, points: 1, grade: true, solution: false -->
```python
l, r = I_x(0, 60, np.pi, 10)
assert l <= np.pi <= r
```

<!-- grade_id: cell-776e14883a36013b, points: 1, grade: true, solution: false -->
```python
l, r = I_x(0, 1, 0.4, 36)
assert l < 0.4 < r
```

Agora, faça o gráfico do erro que os métodos do ponto médio e de Simpson
vão dar quando integrarem o intervalo contendo a "descontinuidade" $\pi$,
em função de $n$.

<!-- grade_id: cell-a667781d9a71248b, points: 0, grade: false, solution: true -->
```python
a, b = 0, 6

# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
```

<!-- grade_id: cell-37ae30958f69fd59, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2
```

<!-- grade_id: cell-22932ced6fcb2767, points: 2, grade: true, solution: false -->
```python
for l in ax.lines:
    assert 1e-10 < min(l.get_ydata()) < 1e-7
```

```python
ax = None
```

## Questão 4: Um intervalo contra todos

Vamos ver quanto do erro de cada método é causado pelo erro em apenas um intervalo.

Faça um gráfico para comparar esses erros:
- à esquerda, os erros para o ponto médio
- à direita, os erros para o método de Simpson

<!-- grade_id: cell-64ea43a87311541c, points: 1, grade: true, solution: true -->
```python
fig, (ax1, ax2) = plt.subplots(ncols=2, figsize=(13,4))

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a27d00cc15e2d685, points: 1, grade: true, solution: false -->
```python
assert len(ax1.lines) == 2
assert len(ax1.legend().texts) == 2
```

<!-- grade_id: cell-b7a54c629922e2b1, points: 1, grade: true, solution: false -->
```python
assert len(ax2.lines) == 2
assert len(ax2.legend().texts) == 2
```

```python
for l in ax1.lines:
    assert 1e-10 < min(l.get_ydata()) < 1e-7
for l in ax2.lines:
    assert 1e-10 < min(l.get_ydata()) < 1e-7
```

```python
ax = None
```

O que você observa quanto aos erros dos métodos?

<!-- grade_id: cell-4770cfe4a90e8a05, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 5: Integração com ajuda humana

Mude a integral de $[0,6]$ por duas integrais, uma em $[0,\pi]$ e outra em $[\pi,6]$.
Faça dois gráficos comparando os erros dos 4 métodos:
- Um para o método do ponto médio, o outro para o método de Simpson
- Em cada um, o erro cometido por uma única integral e o erro cometido pela integral feita em 2 partes

Use o mesmo número $n$ **total** de sub-intervalos, para não "roubar" em favor do que quebra a integral em 2  :-)

<!-- grade_id: cell-089770b7f62c0c00, points: 1, grade: true, solution: true -->
```python
a, c, b = 0, np.pi, 6
exact = F(b) - F(a)

fig, (ax1, ax2) = plt.subplots(ncols=2, figsize=(13,4))
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
```

<!-- grade_id: cell-9a5b67b049806def, points: 1, grade: true, solution: false -->
```python
assert len(ax1.lines) == 2
assert len(ax1.legend().texts) == 2
assert ax1.get_title != ""
```

<!-- grade_id: cell-94ec0f89ceb454d4, points: 1, grade: true, solution: false -->
```python
assert len(ax2.lines) == 2
assert len(ax2.legend().texts) == 2
assert ax2.get_title != ""
```

```python
ax = None
```

O que você observa?
O que você conclui?

<!-- grade_id: cell-9b400f10460b7b4b, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE