![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

# Teste 2: funções inversas

Neste teste, iremos analisar a função $f(x) = x e^x$, para construir a função inversa $g(y) = x$.

```python
import numpy as np
import matplotlib.pyplot as plt
```

# Questão 1: Análise da função

## 1.1 Gráfico

Em geral, uma boa forma de analisar uma função é fazendo o gráfico dela.

Faça um gráfico de $f$ no intervalo $[-4,1]$. Também indique:
- um título; e
- dê nome aos eixos.

<!-- grade_id: cell-b62ae0d9904027d1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-55a02f4e466c71f9, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 1
assert ax.title.get_text() != ""
assert ax.get_xlabel() != ""
assert ax.get_ylabel() != ""
```

<!-- grade_id: cell-5698f335a74400b0, points: 1, grade: true, solution: false -->
```python
valores = ax.lines[0].get_ydata()
a, b = min(valores), max(valores)
assert np.isclose(a, -0.36787944117144, atol=1e-4)
assert np.isclose(b, 2.718281828459045, atol=1e-4)
```

```python
ax = None
```

## 1.2: Mínimos e Máximos

Se uma função troca a direção de crescimento, isso muda a análise das raízes.
Portanto, é importante saber onde estão os máximos e mínimos locais de $f$.

Calcule (analiticamente) a derivada de $f$, e indique o(s) ponto(s) em que ela se anula.

São pontos de máximo ou de mínimo?

<!-- grade_id: cell-afcca791543112ae, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Agora, faça uma figura com dois gráficos:
- No primeiro, o gráfico da derivada de $f$ no mesmo intervalo $[-4, 1]$;
- Ao lado, um "zoom" em uma região que você considere importante.

<!-- grade_id: cell-cdaee38e6b4914a3, points: 2, grade: true, solution: true -->
```python
fig, [ax1, ax2] = plt.subplots(ncols=2, figsize=(12,4))
# YOUR CODE HERE
raise NotImplementedError()
```

Como estes gráficos ajudam a confirmar os mínimos e máximos que você calculou?

<!-- grade_id: cell-2ff292ac34b1722f, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## 1.3: Número de raízes

Usando a informação da derivada da função, seus maximo(s) e minimo(s), responda: 

Quantas raizes a equação $f(x) = y$ tem para cada valor de $y$?

<!-- grade_id: cell-5dde5d2327304877, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

# Questão 2: Construindo a função inversa

A maior dificuldade de usar a bisseção para construir a função inversa é que precisamos de um intervalo onde a raiz estará.

## 2.1: Gráficos para limitar a bisseção

Faça um gráfico das funções $x e^x$ e $e^x$ e veja que,
para $y$ suficientemente grande,
a raiz estará sempre abaixo de $\log y$.
Isso permite achar um limite superior para um intervalo de bisseção.

Escolha um intervalo razoável para o gráfico, que ajude a observar quão "grande" $y$ deve ser.

<!-- grade_id: cell-f37ddc6a9a659c5e, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-90cdc1d79362f65d, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2
assert ax.title.get_text() != ""
```

Agora, falta estudar o que acontece quando $x \to -\infty$.

Faça o gráfico de $f$ e o de $1/x$ para $x \in [-4, -1]$.

<!-- grade_id: cell-0d3b0647ca62048e, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

Explique porque (usando o gráfico e outras contas, se você julgar necessário),
quando $f(x) = y$ tem uma raiz "muito negativa",
esta raiz é menor (em valor absoluto) do que a raiz de $1/x = y$.

<!-- grade_id: cell-29ae2d774355df0e, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## 2.2: A função inversa propriamente dita

Agora, programe a função inversa de $f$ (chamada $g$),
que retorne **todas** as soluções de $f(x) = y$ em uma lista.

Se não houver raízes, retorne a lista vazia.

<!-- grade_id: cell-a187ca48ed58242e, points: 0, grade: false, solution: true -->
```python
def bisseção(f, a, b, xtol=1e-10, ytol=1e-10):
    # Inclua aqui a sua bisseção
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-377500a0b349571b, points: 0, grade: false, solution: true -->
```python
def g(y, xtol=1e-8, ytol=1e-8):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1469ba8d74314f46, points: 1, grade: true, solution: false -->
```python
r1 = g(1)[0]
assert np.isclose(r1, 0.5671432904097838)
assert np.isclose(f(r1), 1)
```

<!-- grade_id: cell-5cd8e074b90655de, points: 1, grade: true, solution: false -->
```python
assert g(-1) == []
```

<!-- grade_id: cell-d061712c94645455, points: 1, grade: true, solution: false -->
```python
r1, r2 = g(-0.1)
assert np.isclose(f(r1), -0.1)
assert np.isclose(f(r2), -0.1)
assert abs(r1 - r2) > 1
```

<!-- grade_id: cell-b9e913444d2545ae, points: 1, grade: true, solution: false -->
```python
r1 = g(200)[0]
assert np.isclose(f(r1), 200)
```

# Questão 3:  Mudança de Variaveis

Vamos modificar um pouco nossa equação $x e^x = y$.

Para cada um dos casos, ache as raízes a partir de uma mudança de variáveis para transformar a equação em algo da forma

$$ u e^u = t. $$
**Escreva sempre qual foi a mudança de variaveis feita**, relacionando $t$ com $y$ e $u$ com $x$.

Dica: Ache uma mudança de variáveis "na mão" para as funções,
utilize o código da questão anterior para descobrir as raízes "transformadas" $u$,
e depois volte para as variáveis $x$

## 3.1: Expoentes diferentes

Resolva $x e^{2x} = 7$

<!-- grade_id: cell-367099a2b6b5b7cf, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-760f980e26af582f, points: 0, grade: false, solution: true -->
```python
# Dê a resposta da forma $x = ...$
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6e5e9b4461a76dfa, points: 1, grade: true, solution: false -->
```python
assert np.isclose(x * np.exp(2*x), 7)
```

## 3.2: Parte linear diferente

Encontre as duas raízes de $(x  + 1) e^{x} = -\frac{1}{10}$.

<!-- grade_id: cell-098e3c114fe9714e, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-4ec726bcfe9735e6, points: 0, grade: false, solution: true -->
```python
# Dê a resposta da forma $x1, x2 = ...$
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-674ff2854bdc9a07, points: 1, grade: true, solution: false -->
```python
assert np.isclose((x1+1)*np.exp(x1), -0.1)
assert np.isclose((x2+1)*np.exp(x2), -0.1)
assert abs(x1 - x2) > 1
```

## 3.3: Mudança genérica de variáveis

Para resolver a equação

$$(x + a)e^{bx} = c, \quad \text{com $b \neq 0$},$$

podemos fazer uma mudança de variáveis (em função de $a$ e $b$)
para transformar em uma equação $u e^u = t$.

Explique qual é a mudança a ser feita.

<!-- grade_id: cell-1e305d2548600ae0, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Agora, faça um programa que recebe os parâmetros $a$, $b$ e $c$
e retorna uma lista com as soluções.

Dica: use uma _list comprehension_  para fazer a transformação inversa.

<!-- grade_id: cell-8e6df9766734c337, points: 0, grade: false, solution: true -->
```python
def solução_geral(a,b,c):
    """Soluções da equação (x+a)e^{bx} = c"""
    assert b != 0
    # YOUR CODE HERE
    raise NotImplementedError()
```

Testes

```python
def geral(x, a, b):
    return (x+a)*np.exp(b*x)
```

<!-- grade_id: cell-f63aaea2b1d1d4ed, points: 1, grade: true, solution: false -->
```python
r1 = solução_geral(0, 1, 2)[0]
assert np.isclose(geral(r1, 0, 1), 2)
```

<!-- grade_id: cell-54a5b058ab3b235c, points: 1, grade: true, solution: false -->
```python
r1 = solução_geral(3, 2, 2)[0]
assert np.isclose(geral(r1, 3, 2), 2)
```

<!-- grade_id: cell-e3f86c2c394c66e2, points: 1, grade: true, solution: false -->
```python
r1 = solução_geral(-1, -1, -1)[0]
assert np.isclose(r1, 0)
```

<!-- grade_id: cell-755d90353c7c3dbd, points: 2, grade: true, solution: false -->
```python
r1, r2 = solução_geral(-4, -1, 0.005)
assert np.isclose(geral(r1, -4, -1), 0.005)
assert np.isclose(geral(r2, -4, -1), 0.005)
assert np.abs(r1 - r2) > 1
```

# Questão 4: Uma outra função... (Bônus)

Vamos aplicar o mesmo raciocínio das questões anteriores para uma função um pouco diferente. 

Faça o gráfico da função $h(x) = x - \log(x)$.

<!-- grade_id: cell-40ba82c82590b4cc, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Ao variar o y, o que acontecerá com o número de raízes?
Essa função pode ter mais de 2 raizes?

<!-- grade_id: cell-8ae773659416e3c0, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE