![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

```python
import numpy as np
import matplotlib.pyplot as plt

from numpy.linalg import norm
```

# Parte 1: Iteração de Rayleigh

Vimos que podemos iterar um vetor $v$ pela matriz $A$, obtendo a sequência de vetores $A^nv$, por multiplicações sucessivas, e que isso permite encontrar um autovetor.

## Questão 1

Implemente uma função `itera(A,v,tol,maxiter)` que itera o vetor $v$, normalizando a cada iteração.

O método para quando o relativo entre $Av_k$ e $\lambda_k v_k$ for menor do que `tol`,
ou após realizar `maxiter` iterações.

A função retorna `(vs, ls)`, respectivamente:
- as estimativas do autovetor (com norma 1);
- as estimativas correspondentes para o autovalor.

<!-- grade_id: cell-257384f421cba661, points: 0, grade: false, solution: true -->
```python
def itera(A,v, tol=1e-6, maxiter=1000):
    v = np.asarray(v)
    A = np.asarray(A)
    n,m = np.shape(A)
    assert n==m, 'A must be square'
    v /= norm(v)
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-80f43e99a64ee409, points: 1, grade: true, solution: false -->
```python
# Autovetores conhecidos
A = [[1,2],[2,1]]
alvo = np.array([1,1])/np.sqrt(2)

vs, ls = itera(A,[1.,2])

assert(abs(ls[-1] - 3) < 1e-6)
assert(all(abs(vs[-1] - alvo) < 1e-6))
assert(10 < len(ls) < 20)
```

<!-- grade_id: cell-abf4f0438eefa28a, points: 1, grade: true, solution: false -->
```python
# Autovetores conhecidos
A = [[1,2],[2,1]]
alvo = np.array([1,1])/np.sqrt(2)

vs, ls = itera(A,[1.,2], tol=1e-12)

assert(abs(ls[-1] - 3) < 1e-12)
assert(all(abs(vs[-1] - alvo) < 1e-12))
assert(20 < len(ls) < 30)
```

## Questão 2: Convergência

Vejamos como o algoritmo "converge" para o autovetor.
Use os vetores intermediários e faça um gráfico da evolução do erro entre os $v$'s produzidos e o autovetor $v_\lambda$.

<!-- grade_id: cell-5fbe094e6250b829, points: 0, grade: false, solution: true -->
```python
ax = None
vs, ls = itera(A, np.random.rand(2), tol=1e-10)

# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-64ee0faf69476729, points: 2, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 1

ys = ax.lines[0].get_ydata()

assert 1e-12 < min(ys) < 1e-10
assert np.all(ys[:-1] > ys[1:])
```

O que o último assert quer dizer?

<!-- grade_id: cell-8432b1317c5a799a, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 3: Convergência comparada

Para cada um dos vetores `d1` e `d2` abaixo, considere a matriz $A = \operatorname{diag}(d_i)$ correspondente.

```python
d1 = [1,10,20,30,31,32]
d2 = [1,10,20,29,30,32]
```

Qual é o autovetor com o maior autovalor para $A_1$ e $A_2$?

<!-- grade_id: cell-80ff3d3111fd3705, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Agora, compare a velocidade de convergência do autovetor usando `itera` para cada uma destas matrizes,
fazendo o gráfico do erro entre os vetores gerados para $A_1$ e $A_2$ no mesmo eixo.

<!-- grade_id: cell-0ca9beaae72f048a, points: 0, grade: false, solution: true -->
```python
ax = None

vs_1, _ = itera(np.diag(d1), np.ones(6))
vs_2, _ = itera(np.diag(d2), np.ones(6))

# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-b5b7998bcab31e94, points: 2, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2
```

Para qual matriz há convergência mais rápida?  Como você explicaria isso?

<!-- grade_id: cell-c2652af3d1fa97e1, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 4: Convergência?

Sejam $\theta \in [0,2\pi]$ e $\alpha \in \mathbb{R}$,
e considere a matriz 
$$A(\theta, \alpha) = \begin{bmatrix} 
\cos(\theta) & \sin(\theta) & 0\\
-\sin(\theta) &\cos(\theta) &  0\\
0 & 0 & \alpha\\
\end{bmatrix}.$$

Qual a interpretação geométrica dessa matriz?

<!-- grade_id: cell-549ef2ace83f1d22, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Quais são os autovetores de $A$ (em função de $\theta$ e $\alpha$)?

<!-- grade_id: cell-c71fb25daa36f1e4, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Implemente a função abaixo que gera a matriz $A$:

<!-- grade_id: cell-782a31adaac5452f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ec8cf18f333dc2d1, points: 1, grade: true, solution: false -->
```python
assert np.allclose(make_matrix(0,1),np.eye(3))
assert np.allclose(make_matrix(np.pi,0.5),[[-1,0,0],[0,-1,0],[0,0,0.5]])
```

Fixando $\theta = \dfrac{\pi}{4}$,
faça um gráfico do número de iterações necessários para calcular o maior autovetor,
em função de $\alpha \in [0.5,1.5]$, com precisão `1e-12`.

<!-- grade_id: cell-b784f741aca3e042, points: 0, grade: false, solution: true -->
```python
alphas = np.linspace(0.5,1.5,100)
ax = None

# YOUR CODE HERE
raise NotImplementedError()


ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-e13220288c2e5c20, points: 2, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 1
assert ax.get_xlabel() != ""

ys = ax.lines[0].get_ydata()

assert 100 > ys.min() > 60
assert ys[55] < 600
assert ys[50] > 900
```

Agora, faça o gráfico com a estimativa do autovalor, novamente em função de $\alpha$.

<!-- grade_id: cell-496cb16ecaa50929, points: 0, grade: false, solution: true -->
```python
ax = None

# YOUR CODE HERE
raise NotImplementedError()


ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-6bc7eba37b6f9533, points: 2, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 1
assert ax.get_xlabel() != ""

ys = ax.lines[0].get_ydata()

assert np.all(0.7 <= ys) and np.all(ys <= 1.5)
```

Como explicar a variação no número de iterações? O que isso tem a ver com o autovalor retornado?

<!-- grade_id: cell-d2c0a1630ea2b95b, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

# Questão 5: Autovalores negativos

Em princípio, tudo deveria funcionar de forma muito similar para autovalores **negativos**.

Estude o caso $\alpha \in [-1.5, -0.5]$, fazendo os dois gráficos lado a lado.

<!-- grade_id: cell-57f58a930c6eee87, points: 0, grade: false, solution: true -->
```python
fig, (ax1, ax2) = plt.subplots(ncols=2, figsize=(13,4))
negalphas = np.linspace(-1.5, -0.5, 100)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-59a7cae1d7d54ba6, points: 2, grade: true, solution: false -->
```python
assert ax1.title.get_text() != ""
assert len(ax1.lines) == 1
assert ax1.get_xlabel() != ""

ys1 = ax1.lines[0].get_ydata()

assert 100 > ys1.min() > 60
assert ys1[44] < 600
assert ys1[49] > 900


assert ax2.title.get_text() != ""
assert len(ax2.lines) == 1
assert ax2.get_xlabel() != ""

ys2 = ax2.lines[0].get_ydata()

assert np.all(0.71 >= ys2)
assert np.all(ys2 >= -1.5)
```

Para funcionar, o teste `ys1[i] < 600` teve seu índice foi trocado de `55` para `44`.
Explique esta mudança, tanto do ponto de vista matemático,
quanto da linguagem Python. 

<!-- grade_id: cell-de42c88dea47bd4a, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Explique, também, porque $0.71$ deve ser suficiente como cota superior. (Aqui, é "só" matemática, mas a matemática pode ser um pouco mais complexa)

<!-- grade_id: cell-cf2020151488e34a, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

# Questão 6: Mudando $\theta$

A velocidade de convergência muda ao trocar $\theta$?
Teste outros valores, e refaça os gráficos, agora com várias curvas em cada eixo,
uma para cada um dos valores de $\theta$ que você estudou.

<!-- grade_id: cell-9b0c0c87e0cb96a7, points: 3, grade: true, solution: true -->
```python
fig, (ax1, ax2) = plt.subplots(ncols=2, figsize=(13,4))

# YOUR CODE HERE
raise NotImplementedError()
```

O que você observa?  Isso faz sentido?

<!-- grade_id: cell-f882f904ce506664, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE