![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

```python
import numpy as np
import matplotlib.pyplot as plt
```

# Parte 1: Cálculo vetorial

## Questão 1: Derivadas direcionais

As derivadas direcionais são obtidas por um limite um pouco mais complicado:

$$ \frac{\partial f}{\partial v}(x)= \lim_{h\to0} \frac{f(x+hv) - f(x)}{h}. $$

(às vezes, também se escreve $\nabla_v f(x)$ ou $f'_v(x)$ para a derivada direcional.)

Generalize a função `df` para que ela calcule derivadas direcionais.

<!-- grade_id: cell-b5f298e9443e23d7, points: 0, grade: false, solution: true -->
```python
def df(f,x,v,h=1e-8):
    # YOUR CODE HERE
    raise NotImplementedError()
```

### Algumas funções vetoriais

```python
def norm1(x):
    return np.sum(np.abs(x))
def norm2(x):
    return np.sum(x**2)
def estranha(x):
    x1,x2 = x
    return np.cos(x1) + 2*np.sin(x2/2)
```

#### Testes simples

<!-- grade_id: cell-248d2b2cb16b7b13, points: 1, grade: true, solution: false -->
```python
assert np.isclose(df(norm2, np.array([3,4]), np.array([0,2])), 16)
assert np.isclose(df(norm2, np.array([3,4]), np.array([1,-1])), -2)
```

<!-- grade_id: cell-99f4fbbaf3e65ab4, points: 1, grade: true, solution: false -->
```python
assert np.isclose(df(estranha, np.array([1,2]), np.array([2,1])), -1.1426397161784507)
assert np.isclose(df(estranha, np.array([1,2]), np.array([2,1]), h=1e-3), -1.1426397161784507, rtol=2e-3)
```

<!-- grade_id: cell-3f795fc1959aa5ee, points: 1, grade: true, solution: false -->
```python
assert np.isclose(df(norm1, np.array([3,3]), np.array([0,2])), 2)
assert np.isclose(df(norm1, np.array([3,3]), np.array([1,-1])), 0)
```

#### Testando propriedades

<!-- grade_id: cell-a2924509bf252fe5, points: 1, grade: true, solution: false -->
```python
assert np.isclose(df(norm2, np.array([0,3]), np.array([1,0])), 
                  -df(norm2, np.array([0,3]), np.array([-1,0])))

assert np.isclose(df(norm1, np.array([0,3]), np.array([1,0])), 
                  df(norm1, np.array([0,3]), np.array([-1,0])))
```

**Pergunta:** Como interpretar (em Cálculo) estas duas últimas igualdades?  Porque a segunda não tem um sinal de menos?

<!-- grade_id: cell-3cc75fd5044d0a46, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 2: Gradientes

Vamos usar a nova função `df` para calcular [gradientes](https://pt.wikipedia.org/wiki/Gradiente) e outros objetos do cálculo vetorial.

Usando a função `len` (para descobrir a dimensão!), implemente `grad(f,x,delta)`,
onde cada derivada parcial é calculada com um passo de tamanho $\delta$.

<!-- grade_id: cell-9e935e4d0b917421, points: 0, grade: false, solution: true -->
```python
def grad(f,x,delta=1e-8):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d2c0bf67345f07ef, points: 1, grade: true, solution: false -->
```python
p = np.array([3,4])
assert np.allclose(grad(norm2, p, delta=1e-5), 2*p, rtol=1e-5)
```

<!-- grade_id: cell-4ac81ccce97336d5, points: 1, grade: true, solution: false -->
```python
assert np.allclose(grad(norm1, np.array([3,4])), [1,1])
assert np.allclose(grad(norm1, np.array([3,-4])), [1,-1])
```

<!-- grade_id: cell-4ff877466a47984b, points: 1, grade: true, solution: false -->
```python
ans = [-0.14112000805986724, -0.41614683654714246]
assert np.allclose(grad(estranha, np.array([3,4]), 1e-8), ans)
```

## Questão 3: Funções vetoriais

Agora, vejamos o que acontece quando a função $f$ vai de $R^n$ em $R^m$.
Supondo que a função (programada) `f` receba um vetor (de dimensão $n$) e retorne um vetor (de dimensão $m$),
dê a fórmula da $j$-ésima coordenada do vetor `df(f,x,v,h)`,
onde `df` é a sua função da questão 1.

**Sugestão:** use $f(p)[j]$ para a $j$-ésima coordenada do vetor `f(p)`.

<!-- grade_id: cell-3a85ad79e947a5f3, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### Mais funções vetoriais

```python
def curva1(t):
    return np.array([np.cos(t), np.sin(t), t])

def superficie1(t):
    u,v = t
    return np.array([u*np.exp(v-u), v*np.cos(u+v), np.sin(v)])
```

<!-- grade_id: cell-221511232c88b938, points: 1, grade: true, solution: false -->
```python
assert np.allclose(df(superficie1, np.array([0,0]), np.array([1,2])), [1,2,2])
```

<!-- grade_id: cell-e77bb4ced51f5cf4, points: 1, grade: true, solution: false -->
```python
ans = [-0.9092974268256819, -0.41614683654714246, 1.0]
assert np.allclose(df(curva1, 2, 1, 1e-5), ans, rtol=2e-5)
```

## Questão 4: Ordem dos eixos

A sua função `grad` deveria retornar um `np.array`, e para a função `superficie1`,
de $R^2$ em $R^3$, isso deve dar a matriz com todas as derivadas parciais.

```python
grad(superficie1, np.array([1,2]), delta=1e-5)
```

Observando o resultado acima, o que você pode dizer sobre as linhas e colunas da matriz gradiente?
Elas coincidem com a ordem típica da matriz Jacobiana do cálculo 2?

<!-- grade_id: cell-6166bb35f1fafef9, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 5: Divergente

Adapte o cálculo do gradiente para obter o divergente de uma função vetorial de $R^n$ em $R^n$:

$$ \text{div} F(x) = \sum \frac{\partial F_i}{\partial x_i}(x). $$

<!-- grade_id: cell-fe75e66e80f7ff82, points: 0, grade: false, solution: true -->
```python
def div(f,x,delta=1e-8):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
def polar(p):
    rho,theta = p
    return rho*np.array([np.cos(theta), np.sin(theta)])

def gravity(p):
    return -p/sum(p**2)
```

<!-- grade_id: cell-dffa4c34454ebdda, points: 1, grade: true, solution: false -->
```python
assert np.isclose(div(polar, np.array([1,0]), delta=1e-3), 2, rtol=1e-3)
```

<!-- grade_id: cell-f9a201d81750a5f5, points: 2, grade: true, solution: false -->
```python
gpolar = grad(polar, np.array([1,0]), delta=1e-3)

assert np.allclose(gpolar, np.eye(2), rtol=1e-3, atol=1e-3)
assert np.sum(np.abs(gpolar - np.eye(2))) > 1e-4
```

<!-- grade_id: cell-eed2eda75a0180c0, points: 1, grade: true, solution: false -->
```python
assert np.isclose(div(gravity, np.array([1,2,1]), delta=1e-6), -1/6, rtol=1e-6)
```

<!-- grade_id: cell-230611dfef1de58c, points: 1, grade: true, solution: false -->
```python
assert np.isclose(div(gravity, np.array([1,1,1,1,1])), -0.6, rtol=1e-8)
```

# Parte 2: Mais precisão

Se, em vez de usarmos derivadas laterais, usarmos derivadas centrais,
deveríamos obter mais precisão na resposta.

## Questão 6: Derivadas centrais

Modifique as funções `grad` e `div` para serem simétricas:

<!-- grade_id: cell-0bfc6d0b8fbfa039, points: 0, grade: false, solution: true -->
```python
def df_sym(f,x,v,h):
    # YOUR CODE HERE
    raise NotImplementedError()

def grad_sym(f,x,delta=1e-5):
    # YOUR CODE HERE
    raise NotImplementedError()

def div_sym(f,x,delta=1e-5):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ba882b706df528b2, points: 1, grade: true, solution: false -->
```python
gpolar = grad_sym(polar, np.array([1,0]), delta=1e-3)

assert np.allclose(gpolar, np.eye(2), rtol=1e-6, atol=1e-6)
assert np.sum(np.abs(gpolar - np.eye(2))) > 1e-8
```

<!-- grade_id: cell-8d96aca4375ddfa1, points: 1, grade: true, solution: false -->
```python
assert np.isclose(div_sym(gravity, np.array([1,2,1])), -1/6, atol=1e-10, rtol=1e-10)
```

## Questão 7: A derivada é linear

Sabemos que, se a função $F : R^n \to R^m$ for derivável,
então podemos calcular derivadas direcionais usando a matriz jacobiana $DF(x)$:
$$
  \frac{\partial F}{\partial v}(x) = DF(x) \cdot v.
$$

```python
x = np.array([1,2])
vs = [np.array([1,3]), np.array([3,1]), np.array([1,1]), np.array([1,-1])]

for v in vs:
    DF_norm2 = grad(norm2, x, delta=1e-3)
    parcial  = df(norm2, x, v, h=1e-3)
    print(DF_norm2 @ v - parcial)
```

Porquê aparece este erro?

<!-- grade_id: cell-3fefd1494a881def, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Expanda em série de Taylor a fórmula da derivada lateral e a fórmula do "gradiente lateral",
e encontre uma expressão para o erro entre
1. a derivada lateral; e
2. a matriz jacobiana vezes o vetor diretor.

<!-- grade_id: cell-20a2721b6135ff86, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 8: melhorando a linearidade?

Agora, vejamos o que acontece com derivadas simétricas:

```python
for v in vs:
    DF_norm2 = grad_sym(norm2, x, delta=1e-3)
    parcial  = df_sym(norm2, x, v, h=1e-3)
    print(DF_norm2 @ v - parcial)

print()
for v in vs:
    DF_norm2 = grad_sym(estranha, x, delta=1e-3)
    parcial  = df_sym(estranha, x, v, h=1e-3)
    print(DF_norm2 @ v - parcial)
```

O erro é menor, mas não apenas.
O que acontece de especial no caso da norma 2?

Dica: expanda a série de Taylor para ver qual o erro da derivada central.

<!-- grade_id: cell-3a505f0c8892c673, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE