![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

```python
import numpy as np
import matplotlib.pyplot as plt
```

# Teste 3: Método de Newton

Vamos continuar o problema de calcular a função inversa de $f(x) = x e^x$.

# Questão 1: Newton para inversão

Vamos generalizar o método de Newton para encontrar raizes de $f(x) = y$.

Inclua abaixo tanto o método de Newton, retornando **todos** os pontos percorridos até encontrar a raiz.

<!-- grade_id: cell-01488b62ae1c6a0c, points: 0, grade: false, solution: true -->
```python
def newton_y(f, df, y, x0, xtol=1e-12, maxiter=50):
    """Calcula uma raiz aproximada de  f(x) = y  pelo método de Newton, a partir do ponto x0."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e0ebc41b3dd240e9, points: 1, grade: true, solution: false -->
```python
def f1(x):
    return x * np.log(x)

def df1(x):
    return np.log(x) + 1

pts_newton = newton_y(f1, df1, 1.1, 1)
assert abs(f1(pts_newton[-1]) - 1.1) < 1e-10
```

<!-- grade_id: cell-a3f1b983ac6e7995, points: 1, grade: true, solution: false -->
```python
pts_newton = newton_y(f1, df1, 1.1, 1)

assert abs(f1(pts_newton[1]) - 1.1) > 1e-3
assert 3 < len(pts_newton) < 10
```

<!-- grade_id: cell-d0aeb5d95ed08e07, points: 1, grade: true, solution: false -->
```python
def f2(x):
    return x * np.sin(x)

def df2(x):
    return np.sin(x) + x * np.cos(x)

pts_newton = newton_y(f2, df2, 1.1, 1)
assert abs(f2(pts_newton[-1]) - 1.1) < 1e-10
```

# Questão 2: Gráficos

O método de Newton funciona bem quando a primeira derivada está suficientemente longe de zero.

## Questão 2.1: Derivada
Faça um gráfico da derivada $f'(x)$, e deduza as regiões onde o método de Newton provavelmente funciona bem,
e onde ele funciona mal.

```python
def f(x):
    return x * np.exp(x)

def df(x):
    return (x+1) * np.exp(x)
```

<!-- grade_id: cell-591a25b1a2d849ac, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-03b0cbc40cda2de8, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) >= 1
assert len(ax.title.get_text()) > 0
```

```python
ax = None
```

Em quais regiões o gráfico sugere que será difícil calcular uma raiz de $f(x) = y$?

<!-- grade_id: cell-3d08a4024fc168c6, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 2.2: Número de iterações para Newton

Faça o gráfico do número de iterações necessárias para o método de Newton encontrar uma raiz de $f(x) = y$
para $y$ no intervalo $[1,100]$, e partindo de $x_0 = 1$.

<!-- grade_id: cell-aabddeab26b8e124, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-0e4640465a4c1e5a, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 1
assert len(ax.title.get_text()) > 0
```

<!-- grade_id: cell-8ee2f33778a564c2, points: 1, grade: true, solution: false -->
```python
assert len(ax.get_xlabel()) > 0
assert len(ax.get_ylabel()) > 0
```

<!-- grade_id: cell-62d3bf6b9b9f6da9, points: 1, grade: true, solution: false -->
```python
ydata = ax.lines[0].get_ydata()
assert 5 <= min(ydata) <= 7
assert 22 <= max(ydata) <= 30
```

```python
ax = None
```

Agora, para cada $y$ (no mesmo intervalo), faça o gráfico
- da raiz encontrada; e
- do maior ponto testado pelo método de Newton,

no mesmo eixo.

<!-- grade_id: cell-06174d781ae4f044, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-712cf22d5d02c889, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 2
assert len(ax.get_legend().texts) == 2
assert len(ax.title.get_text()) > 0
```

```python
ax = None
```

Como este gráfico ajuda a explicar o anterior?

<!-- grade_id: cell-96e21b6497f70130, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 2.3: Outro ponto inicial

Agora, considere que o ponto inicial para encontrar uma raiz de $x e^x = y$ é $\log y$.
Quantas iterações são necessárias?

<!-- grade_id: cell-06b2958864fdff69, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-77910f3b72a808c6, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 1
assert len(ax.title.get_text()) > 0
```

```python
ax = None
```

O que aconteceu?

<!-- grade_id: cell-47fc3651761a578b, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Refaça os gráficos do maior ponto e da raiz.

<!-- grade_id: cell-aa72ed3f225df84c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

Como este gráfico ajuda a explicar os anteriores?

<!-- grade_id: cell-5ee02b5d89dd29fa, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

# Questão 3: Regiões difíceis

Vamos agora observar quão difícil é calcular uma raiz de $x e^x = y$ para pontos onde a derivada é praticamente nula.

## Questão 3.1: Ponto de mínimo

Faça um gráfico do número de iterações do método de Newton para calcular a raiz de $x e^x = y$ para $y$ no intervalo $[-1/e, 0]$, começando a partir de $x = 0$.

<!-- grade_id: cell-bb6e6b4570d28887, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-ff67a4e8c14669e4, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 1
assert len(ax.title.get_text()) > 0
```

<!-- grade_id: cell-8d455fd0f45ea965, points: 1, grade: true, solution: false -->
```python
vs = ax.lines[0].get_ydata()
assert max(vs) > 20
assert min(vs) < 5
```

```python
ax = None
```

O que você observa no gráfico?

<!-- grade_id: cell-0dfc13abccdfb67b, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Agora, dê um zoom na parte que dá mais iterações.

Dica: o estilo `".:"` para a linha pode ajudar a visualizar o que está acontecendo.

<!-- grade_id: cell-d0e3e8cf1aaf3f6a, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-076cdce5a3e2ac7a, points: 2, grade: true, solution: false -->
```python
vs = ax.lines[0].get_ydata()
assert max(vs) > 20
assert 10 < min(vs) < 15
```

```python
ax = None
```

## Questão 3.2: Erros

Agora, vamos construir $y = z e^z$ para $z$ entre 0 e -1.
Assim, já saberemos qual a raiz de $x e^x = y$.

Com isto, faça o gráfico do **erro absoluto** entre a raiz encontrada pelo método de Newton - começando de 0, como na questão anterior - e o valor real.

<!-- grade_id: cell-0050c46eea198121, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

## Questão 3.3: Um outro ponto de vista

Vamos fazer um zoom na parte que dá maior erro.

Para poder observar bem o que acontece, vamos olhar para $z = -1 + w$, e fazer o gráfico do erro em função de $w$.

Dica: para escolher a melhor escala para os eixos, teste as diferentes combinações!

<!-- grade_id: cell-8d08e83c18253fbc, points: 2, grade: true, solution: true -->
```python
ws = np.logspace(-53,0, num=200, base=2)
zs = -1 + ws
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

O que você observa?

<!-- grade_id: cell-3b26dd31709ef8ce, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE