![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_ , depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

# Teste 0

Este teste servirá como "teste" da plataforma de testes e provas para o curso de Computação Científica I.
Assim, você poderá ter uma ideia de como serão corrigidas
- as listas de exercícios
- as provas

## Questão 1

Escreva uma função `cubes` que retorna a lista de todos os cubos da lista passada em argumento.

<!-- grade_id: cell-f0b9d04bcf0176e9, points: 0, grade: false, solution: true -->
```python
def cubes(l):
    """ Essa função retorna uma lista dos cubos dos elementos da lista. """
    ls = [x**3 for x in l]
    return ls
```

```python
# Teste aqui a sua função antes de passar para a parte dos "asserts" abaixo.
cubes([1,2,5]) # Deveria produzir a lista [1, 8, 125]
```

Output:
```
[1, 8, 125]
```

A caixa a seguir contém testes que serão efetuados sobre a sua função.
- Se algum teste falhar, você verá uma mensagem com `Assertion Error`.
- Senão, você verá a mensagem "Tudo ok". (por causa do "print"; em geral, será apenas vazio como nas seguintes)

Não apague: é ela que conta os pontos da questão!

<!-- grade_id: cubes_test, points: 3, grade: true, solution: false -->
```python
assert(cubes([0]) == [0])
assert(cubes([1,2,5]) == [1, 8, 125])
assert(cubes([3, -10, 11]) == [27, -1000, 1331])

print("Tudo ok")
```

Output:
```
Tudo ok
```

**Observação**:
Você pode criar novas caixas no notebook, em especial para testar suas funções.
Isso é útil, em particular, porque as caixas de `assert` não devem ser alteradas
(e o notebook protege para não alterar).

<!-- grade_id: cell-dab4951ec2dfe199, points: 1, grade: true, solution: false -->
```python
# Mais um teste :-)
assert(cubes([]) == [])
```

## Questão 2

Escreva uma função que calcula o máximo das diferenças consecutivas entre os elementos de uma lista.
Em "fórmulas":

$$ g(l) = \max_{i} (l_i - l_{i-1}). $$

<!-- grade_id: cell-ffb7646c93c15d6d, points: 0, grade: false, solution: true -->
```python
def g(l):
    dif = [ l[i] - l[i-1]  for i in range(1, len(l)) ]
    return max(dif)
```

<!-- grade_id: cell-a459c206ee62e94b, points: 3, grade: true, solution: false -->
```python
assert g([0,1,2,3,4,5,6]) == 1
assert g([0,2,3,-1,0]) == 2
assert g([6,5,4,3,2,1,0]) == -1
```

A caixa a seguir também contém testes, mas agora eles estão "escondidos".

Vale lembrar, mais uma vez, a mesma observação de antes: mantenha esta caixa: ela contém os pontos referentes aos testes, mesmo que você não os possa ver.

<!-- grade_id: cell-558c38603d84d3ca, points: 2, grade: true, solution: false -->
```python
# Caixa com testes - não apague
```

### Observação

Para ter certeza que este notebook será corrigido corretamente, o melhor é rodar tudo "de uma vez só, do zero".
Para simular isto, use o menu: `Kernel -> Restart & Run All`.