![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

```python
import numpy as np
```

# Teste 1: Bisseção

## Questão 1: Número de iterações da bisseção, e de chamadas à função `f`

Generalize o algoritmo da bisseção _iterativa_ para retornar,
- a (aproximação da) raiz,
- o número de bisseções (divisões por 2 do intervalo) feitas,
- o número de vezes que você chamou a função `f`

Use como critério de parada
- `xtol`, e
- `maxiter`,

ou seja, o algoritmo termina quando:
- seja possível garantir que o erro (absoluto) da resposta ("em $x$") seja menor do que `xtol`; OU
- quando o algoritmo já tiver feito `maxiter` itererações.

<!-- grade_id: cell-9c2f0bf1fe48da91, points: 0, grade: false, solution: true -->
```python
def bissecao(f, a, b, xtol=1e-8, maxiter=10):
    a,b = sorted( [a,b] )
    
    err = abs(b-a)
    aux = a
    it = 0
    ncalls = 0 

    while err >= xtol :
        if (it > maxiter):
          return x, it-1, ncalls

        x = (a+b)/2.0
        
        if f(x)*f(a) < 0:
            b = x
            ncalls +=2
        else:
            ncalls +=2
            if f(x)*f(a) == 0:
                return x, it, ncalls
            a = x


        err = abs(x-aux)
        aux = x
        it += 1

    return x, it, ncalls
```

### 1.1: Testando com $\sqrt2$:

```python
def p2(x): return x**2 - 2
```

<!-- grade_id: cell-6216e0119f2a0e78, points: 1, grade: true, solution: false -->
```python
raiz, niters, ncalls = bissecao(p2, 0, 2, maxiter=100)
assert abs(raiz - np.sqrt(2)) < 1e-8
assert 25 <= niters <= 30
assert 25 <= ncalls
```

<!-- grade_id: cell-caee585cfc6b637f, points: 1, grade: true, solution: false -->
```python
raiz, niters, ncalls = bissecao(p2, 0, 2, xtol=1e-3)
assert abs(raiz - np.sqrt(2)) < 2e-3
assert 9 <= niters <= 11
```

<!-- grade_id: cell-ed17d59cca42278d, points: 1, grade: true, solution: false -->
```python
raiz, niters, ncalls = bissecao(p2, 0, 2)
assert abs(raiz**2 - 2) > 1e-3
assert niters == 10
```

<!-- grade_id: cell-19379484bc5a3adf, points: 1, grade: true, solution: false -->
```python

```

### 1.2: Tolerância, número de iterações, respostas, ...

Observe o seguinte código:

```python
for err in np.logspace(-8, 0, num=9):
    r, ni, nc = bissecao(p2, 0, 2, xtol=err)
    print(f"xtol = {err:.1e} --> raiz = {r:.10f}, #iter = {ni:3d}")
```

Output:
```
xtol = 1.0e-08 --> raiz = 1.4150390625, #iter =  10
xtol = 1.0e-07 --> raiz = 1.4150390625, #iter =  10
xtol = 1.0e-06 --> raiz = 1.4150390625, #iter =  10
xtol = 1.0e-05 --> raiz = 1.4150390625, #iter =  10
xtol = 1.0e-04 --> raiz = 1.4150390625, #iter =  10
xtol = 1.0e-03 --> raiz = 1.4150390625, #iter =  11
xtol = 1.0e-02 --> raiz = 1.4140625000, #iter =   8
xtol = 1.0e-01 --> raiz = 1.4375000000, #iter =   5
xtol = 1.0e+00 --> raiz = 1.5000000000, #iter =   2
```

Comente

<!-- grade_id: cell-856d87a71e8f8ec5, points: 2, grade: true, solution: true -->
Nesse código, é feito um loop de acordo com o log de n para definicao do erro, utilizando assim, em cada iteracao do for, um valor diferente como criterio de parada para o valor da tolerancia xtol.

Como seria o equivalente se o critério fosse o número de iterações?

<!-- grade_id: cell-c6d45f6ca0f539e4, points: 2, grade: true, solution: true -->
```python
for i in range(0,10):
    r, ni, nc = bissecao(p2, 0, 2, xtol=1e-8, maxiter= i)
    print(f"xtol = {err:.1e} --> raiz = {r:.10f}, #iter = {ni:3d}")
```

Output:
```
xtol = 1.0e+00 --> raiz = 1.0000000000, #iter =   0
xtol = 1.0e+00 --> raiz = 1.5000000000, #iter =   1
xtol = 1.0e+00 --> raiz = 1.2500000000, #iter =   2
xtol = 1.0e+00 --> raiz = 1.3750000000, #iter =   3
xtol = 1.0e+00 --> raiz = 1.4375000000, #iter =   4
xtol = 1.0e+00 --> raiz = 1.4062500000, #iter =   5
xtol = 1.0e+00 --> raiz = 1.4218750000, #iter =   6
xtol = 1.0e+00 --> raiz = 1.4140625000, #iter =   7
xtol = 1.0e+00 --> raiz = 1.4179687500, #iter =   8
xtol = 1.0e+00 --> raiz = 1.4160156250, #iter =   9
```

Comente, também.

<!-- grade_id: cell-b7931432898e9741, points: 2, grade: true, solution: true -->
Nesse caso agora, o loop foi feito em cima do valor do numero de iteracoes, utilizando como criterio de parada o valor que está variavel dentro do for de i, que representa o maximo do numero de iteracoes em cada chamada da funcao de bissecao.

### 1.3: Agora, um polinômio um pouco mais complicado:

```python
def p4(x): return x**4 - 3*x + 1
```

<!-- grade_id: cell-b8f933e77c525f69, points: 1, grade: true, solution: false -->
```python
r1, n1, c1 = bissecao(p4, 0, 1, maxiter=100)
assert p4(r1) < 1e-6
```

<!-- grade_id: cell-fab8ee9c8e55f544, points: 1, grade: true, solution: false -->
```python
r2, n2, c2 = bissecao(p4, 1, 2, maxiter=100)
assert n1 == n2
assert p4(r2) < 1e-6
```

<!-- grade_id: cell-faf0c1fe3a384747, points: 1, grade: true, solution: false -->
```python
r3, n3, c3 = bissecao(p4, 1, 3, maxiter=100)
assert n3 == n2+1
assert r3 == r2
```

<!-- grade_id: cell-1ec4a98add6ac67a, points: 1, grade: true, solution: false -->
```python

```

### 1.4: Relação entre `ncalls` e `niters`

Muito provavelmente, é possível calcular `ncalls` a partir de `niters`.
Qual a relação, para o seu programa, entre estes valores?
Porquê?

<!-- grade_id: cell-4a9c83444fc8ee6a, points: 1, grade: true, solution: true -->
```python
print(n1,c1)
print(n2,c2)
print(n3,c3)
```

Output:
```
27 54
27 54
28 56
```

<!-- grade_id: cell-7d3444498035a363, points: 3, grade: true, solution: true -->
É possível reparar que o numero de chamadas a funcao $f$ é o dobro do número de iterações do programa de bisseção escrito. Isso acontece porque a cada itereção é feita a comparacao entre os valores de $f(x)*f(a)$ para verificação se é menor ou maior que $0$. O objetivo é que seja igual a $0$, então caso seja, o programa encerra, mas caso nao seja, prossegue com a iteração. Portanto a cada iteração, a função $f$ é chamada duas vezes para a forma que o código foi escrito.

## Questão 2: Bisseção em listas ordenadas

Se desejamos encontrar um número em uma lista, podemos percorrer cada um de seus índices
até achar o número, e retornamos o índice correspondente.

Mas, se a lista `l` possui números em ordem,
é possível usar uma variante da bisseção para encontrar este índice.
Observe que agora desejamos retornar um número **inteiro** (o índice!),
e que índices para listas sempre devem ser inteiros, portanto:
- cuidado ao escolher o "ponto médio"
- mas o critério de parada é "quando achar o índice".

### 2.1: Encontrando um número

Implemente a função `bissect_list(l, v)`, que retorna
- o índice `k` do elemento na lista `l`, que vale `v`; e
- o número de acessos à lista `l`.

Suponha que a lista `l` é não-decrescente (ou seja, `l[i] <= l[i+1]`).

<!-- grade_id: cell-8c9f49b0ba8b5d04, points: 0, grade: false, solution: true -->
```python
def bissect_list(l, v):
    if (len(l)== 0):
      return None, 0

    a, b = 0, len(l)-1
    
    k = a
    ncalls = 0 

    if l[b] == v:
      ncalls +=1
      return b, ncalls

    while l[k] != v:
      
      temp = k

      aux_k = (a+b)/2
      k = int(k) #transformando
      
      if (aux_k != k):
        k+=1

      if (k==temp):
        return None, ncalls
      

      if ( (l[a]-v) > (l[b]-v) ):
        a = k
      else: 
        b = k
      
      ncalls +=3
      
    return k, ncalls
      
```

<!-- grade_id: cell-d0e135df902b3955, points: 1, grade: true, solution: false -->
```python
assert bissect_list([], 42) == (None,0)
```

<!-- grade_id: cell-aa0a59b135a65159, points: 1, grade: true, solution: false -->
```python
l = [4]
assert bissect_list(l, 4) == (0,1)
```

<!-- grade_id: cell-ba03109833002639, points: 1, grade: true, solution: false -->
```python
assert bissect_list(l, 5)[0] is None
```

<!-- grade_id: cell-84b5ba3a3384d223, points: 1, grade: true, solution: false -->
```python
l = [1,3,6,10,15,21,28]
for i,v in enumerate(l):
    idx, n = bissect_list(l, v)
    assert idx == i, f"v={v}, i={i}, idx={idx}"
    print(n)
```

Output:
```
0
3
6
9
12
15
1
```

```python
l = [28,21,15,10,6,3,1]
for i,v in enumerate(l):
    idx, n = bissect_list(l, v)
    assert idx == i, f"v={v}, i={i}, idx={idx}"
    print(n)
```

Output:
```
0
3
6
9
12
15
1
```

Qual o maior número de acessos à lista que você precisou para encontrar um valor nela?
Este número depende da posição na lista?  Explique!

<!-- grade_id: cell-9059e7d02c8a0881, points: 2, grade: true, solution: true -->
Depende sim, para os primeiros valores da lista, é necessario um menor numero de chamadas a lista $l$, apenas para o ultimo valor da lista, que é obtido em apenas em uma chamada já que foi implementado um condicional para este caso.

### 2.2: Uma lista muito maior

```python
l = np.arange(-4000, 6000)
np.random.seed(1)
valores_testes = np.random.randint(low=-4000, high=6000, size=10)
for v in valores_testes:
    idx, n = bissect_list(l, v)
    assert l[idx] == v
    #assert n <= 20
    print(f"v={v:5d}, idx={idx:5d}, n={n:3d}")
```

Output:
```
v=-3765, idx=  235, n=705
v= 1192, idx= 5192, n=15576
v=-3095, idx=  905, n=2715
v= 3813, idx= 7813, n=23439
v=-1105, idx= 2895, n=8685
v= 1056, idx= 5056, n=15168
v=-3856, idx=  144, n=432
v=  225, idx= 4225, n=12675
v= 3751, idx= 7751, n=23253
v= -538, idx= 3462, n=10386
```

Como o número de acessos varia, nesta outra lista?
E como ele se compara com o tamanho da lista?

<!-- grade_id: cell-1853e505f297b9b5, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### 3.3: Listas crescentes e decrescentes.

Altere o código anterior, para receber um argumento a mais, `decr`, que indica que a lista é **decrescente**.
Não reordene a lista!

<!-- grade_id: cell-e442fe23480d4b8d, points: 0, grade: false, solution: true -->
```python
def bissect(l, v, decr=False):

    if (len(l)== 0):
      return None, 0

    a, b = 0, len(l)-1
    
    k = a
    ncalls = 0 

    if l[b] == v:
      ncalls +=1
      return b, ncalls

    while l[k] != v:
      
      temp = k

      aux_k = (a+b)/2
      k = int(k) #transformando
      
      if (aux_k != k):
        k+=1

      if (k==temp):
        return None, ncalls
      
      if ( (l[a]-v) > (l[b]-v) ):
        a = k
      else: 
        b = k
      
      ncalls +=3
      
    return k, ncalls
      
```

<!-- grade_id: cell-8d0e1c185171dbce, points: 1, grade: true, solution: false -->
```python
l = [1,3,6,10,15,21,28,35,36]
assert bissect(l, 15)[0] == 4
```

<!-- grade_id: cell-75e5e557da5782e9, points: 1, grade: true, solution: false -->
```python
l_rev = [36, 35, 28, 21, 15, 10, 6, 3, 1]
assert bissect(l_rev, 15, decr=True)[0] == 4
```

<!-- grade_id: cell-bd452b4939210a9a, points: 1, grade: true, solution: false -->
```python
l_rev = [36, 35, 28, 21, 15, 10, 6, 3, 1]
assert bissect(l_rev, 16, decr=True)[0] == None
```

<!-- grade_id: cell-4f806045de0b056c, points: 1, grade: true, solution: false -->
```python
l3 = [36, 35, 28, 21, 15, 10, 9, 6, 3, 2, 1]
assert bissect(l3, 15, decr=True)[0] == 4
```

### 2.4: Bônus

O número de acessos à lista para encontrar um número muda se esta está invertida?
Explique!

<!-- grade_id: cell-def3480b5f0dcea3, points: 2, grade: true, solution: true -->
Não nesse caso, já que para o código desenvolvido,os indices do intervalo são analisados a fim de diminuir o intervalo da nova iteracao. Desse modo,o valor desejado $v$ é comparado aos valores da lista correspondente aos indices do limite do intervalo $(a,b)$, e aquele que possui a maior diferença, será aquele substituido por um novo indice $k$ para nova iteração, até que se obtenha apenas o valor restante.

---

